# CMPS 140

```python
class Node:
    """
    This class describes a single node contained within a graph. 
    It has the following instannce level attributes:
    
    ID: An integer id for the node i.e. 1
    heuristic_cost: A float value representing the estimated 
                    cost to the goal node
    """    
    def __init__(self, ID, heuristic_cost):
        self.ID = ID
        self.connected_nodes = []
        self.heuristic_cost = heuristic_cost
        
    def __repr__(self):
        ID = self.ID
        hx = self.heuristic_cost
        if len(self.connected_nodes)==0:
            nodes = 'None'
        else:
            nodes = ','.join(str(cn[1].ID) for cn in self.connected_nodes)
        return 'Node:{}\nh(n):{}\nConnected Nodes:{}'.format(ID, hx, nodes)
        
    def set_connected_nodes(self,connected_nodes):
        """
        Adds edges that lead from this node to other nodes:
        
        Parameters:
        - connected_nodes: A list of tuples consisting of (cost, Node), 
                           where 'cost' is a floating point value 
                           indicating the cost to get from this node 
                           to 'Node' and 'Node' is a Node object
        """
        self.connected_nodes = connected_nodes
    
def build_graph():
    """
    Builds the graph to be parsed by the search algorithms.
    Returns: The starting node, which is the entry point into the graph
    """
    ids = range(13)
    coords = [(0,0), (1,1), (1,0), (1,1), (5,2), (3,1), (3,0), 
              (3,-1), (5,1), (4,1), (4,0), (4,-2), (7,0)]
    
    #https://en.wikipedia.org/wiki/Euclidean_distance
    euclidean_distance = lambda x1y1, x2y2: ((x1y1[0]-x2y2[0])**2 +  (x1y1[1]-x2y2[1])**2)**(0.5)
    
    def build_connected_node_list(from_id, to_ids):
        starting_coords = coords[from_id]
        
        connected_nodes = []
        for to_id in to_ids:
            connected_nodes.append((euclidean_distance(starting_coords, coords[to_id]), all_nodes[to_id]))
            
        return connected_nodes
    
    goal_coords = (7,0)
    all_nodes = [Node(_id, euclidean_distance(coord, goal_coords)) for _id, coord in zip(ids, coords)]
    
    all_nodes[8].set_connected_nodes(build_connected_node_list(8, [12]))
    all_nodes[10].set_connected_nodes(build_connected_node_list(10,[12]))
    all_nodes[5].set_connected_nodes(build_connected_node_list(5, [8]))
    all_nodes[6].set_connected_nodes(build_connected_node_list(6, [9, 10]))
    all_nodes[7].set_connected_nodes(build_connected_node_list(7, [11]))
    all_nodes[1].set_connected_nodes(build_connected_node_list(1, [4,5]))
    all_nodes[2].set_connected_nodes(build_connected_node_list(2, [5,6]))
    all_nodes[3].set_connected_nodes(build_connected_node_list(3, [7]))
    all_nodes[0].set_connected_nodes(build_connected_node_list(0, [1,2,3]))
    
    return all_nodes[0]
```

```python
# The starting node. You can use this cell to familiarize
# yourself with the node/graph structure
build_graph()
```

Output:
```
Node:0
h(n):7.0
Connected Nodes:1,2,3
```

## Question 1

![graph](https://docs.google.com/drawings/d/e/2PACX-1vStncj9Nc0LddQeViaYnykNxEZsJoYJMHhub2LLX8s7k7gwYjlnlt0cCcivymFihiZyOOMtHwzk1Z4G/pub?w=960&amp;h=720)

Given the above graph implement the following search algorithms:

1. **Depth First Search** (Graph Style) (Textbook Section 3.4.3)
2. **Iterative Deepening Depth-First Search** (Textbook Section 3.4.5)
3. **A* Search** (A-Star) (Textbook Section 3.5.2)



You will implement each of these search algorithms as a function with the following signature:

```python
def search_algorithm(starting_node, goal_node):
    """
    This function implements a search algorithm.
    Parameters:
    - starting_node: The entry node into the graph
    - goal_node: The integer ID of the goal node.
    
    Returns:
    A list containing the visited node ids in order they were visited with starting node
    always being the first node and the goal node always being the last
    """
    
    visited_nodes_in_order = []
    
    # algorithm implemented here
    
    return visited_nodes_in_order
```


I have put skeleton methods for each of the algorithms for you to fill in.

Here are the pseudo code/descriptions implemenatations from the book.

# Depth First Search
---
![dfs.png](https://lh3.googleusercontent.com/y8b0v4eTG2xfUqkWFm7H_RTb_sJ1ERBiyfWjkClb7GtRUsbAFwpECMnL9JfyTcY4AAhBhCWQiQlLhAmWPQLAN_6Ja_WMFEn5U0PoXaPqNJvKjEf9hW6Meq5xyHdygXqB-WzdBRmVEA=w2400)
![dfs-desc.png](https://lh3.googleusercontent.com/dIDKkPQZ7Vn-HfQ0WPlpthT_RkrQ86lH11hjLceUcDvFF2BCZs2RrfRfEiXgpZVG35vkGkLg6q3nISYeEUriPfwVxSNmhxOO0Gujl43G1lRsN8dTBTaxT8baX8tMu0st49R_CsWBKw=w2400)

# Iterative Deepening Depth First Search
---
![iddfs.png](https://lh3.googleusercontent.com/HrkCyMnoMPZoDELRdqx1cakdgA2rPZ2ImIg5n5XE-A5MXwHch7ZZ6uBE4vmVXv34V6BRFOgyrHl6aTXMm0qijA8s47zda7s9dKDDYzEH1mTnhEPWvOM9r1gE1wwXfXP0_3T22zFO0Q=w2400)
![dls.png](https://lh3.googleusercontent.com/w78lEVXUJitCh8WebB49DACKztyErtXTPi4ybHE0ai5L1orFkdYMtc7AfzdiXF1uuotahgqXoYHeaBpeExztTFR0GvKMIE4ddrYoHe6qSea3zL9TQh6LRYHV69cA3pXGt1nyl3Injg=w2400)

# A* Search
---
![astar](https://lh3.googleusercontent.com/Mp31KIl0a7Eu3ky0WU2QaqMBaAZRFH2ed48ua0dAr-P7GQqwyyYi7MT0NzP-2yaF-YQaykasARoUUUolghYc_Vszd4JCYSbKHQMUaAiT2Np9Khypz7NhwMhDOybJGJPdb1e_6oU1RQ=w2400)

# NOTES
---
 - If two nodes are considered equally good chocices, take the node with the lower ID first. 
 - Do not revisit already explored nodes when implementing the Depth First Search.
 
 

<!-- grade_id: cell-8782a75d513b1327, points: 0, grade: false, solution: true -->
```python
DEBUG = False
def dls(start, depth, goal, vn):
    vn.append(start.ID)
    if start.ID == goal:
        return True, vn
    elif depth <= 0:
        return False, vn

    for node in start.connected_nodes:
        r, vn = dls(node[1], depth - 1, goal, vn)
        if r is True:
            return True, vn

    return False, vn


def iterative_deepening_depth_first_search(starting_node, goal_node):
    visited_nodes_in_order = [starting_node.ID]
    MAX_DEPTH = 2^30
    for i in range(1, MAX_DEPTH):
        result, visited_nodes_in_order = dls(starting_node, i, goal_node, visited_nodes_in_order)
        if result is True:
            if DEBUG:
                print(visited_nodes_in_order)
            return visited_nodes_in_order


   
def depth_first_search(starting_node, goal_node):
    frontier, explored = starting_node.connected_nodes[::-1], [starting_node.ID]
    
    if frontier is None:
        return -1
    if starting_node.ID is goal_node:
        return explored
    
    while True:
        if len(frontier) > 0:
            node = frontier.pop()
            node = node[1]
            explored.append(node.ID)
              
            if node.ID == goal_node:
                if DEBUG:
                    print(explored)
                return explored

            for n in (node.connected_nodes[::-1]):
                if n not in explored:
                    frontier.append(n)
        else:
            raise Exception("Empty frontier")
      

    
def a_star_search(starting_node, goal_node):
    closed_nodes = []
    open_nodes = set([(0,starting_node)])
    g = {}
    f = {(0,starting_node):0}
    came_from = {}
    g[(0,starting_node)] = 0

    while len(open_nodes) > 0:
        current_f, current_node = None, None
        for node in open_nodes:
            if current_node is None or f[node] < current_f:
                current_f = f[node]
                current_node = node

        if current_node[1].ID == goal_node:
            taken_path = [current_node[1].ID]

            while current_node in came_from:
                current_node = came_from[current_node]
                taken_path.append(current_node[1].ID)
            taken_path.reverse()
            if DEBUG:
                print(taken_path)
            return taken_path

        closed_nodes.append(current_node)
        open_nodes.remove(current_node)

        for n in current_node[1].connected_nodes:
            if n[1] in closed_nodes:
                continue
            candidate = g[current_node] + n[0]
            if n not in open_nodes:
                open_nodes.add(n)
            elif candidate >= g[n]:
                continue

            came_from[n] = current_node
            g[n] = candidate
            f[n] = candidate + f[current_node]

    return []
```

<!-- grade_id: cell-0c19db32d84a442e, points: 0.0, grade: true, solution: false -->
```python
goal_node = 12
depth_first_search_answer = [0, 1, 4, 5, 8, 12]
iterative_deepening_depth_first_search_answer = [0, 0, 1, 2, 3, 0, 1,
                                                 4, 5, 2, 5, 6, 3, 7,
                                                 0, 1, 4, 5, 8, 2, 5,
                                                 8, 6, 9, 10, 3, 7, 11,
                                                 0, 1, 4, 5, 8, 12]
a_star_search_answer = [0, 2, 6, 10, 12]

assert (depth_first_search(build_graph(), goal_node)==depth_first_search_answer)
assert (iterative_deepening_depth_first_search(build_graph(), goal_node)==iterative_deepening_depth_first_search_answer)
assert (a_star_search(build_graph(), goal_node)==a_star_search_answer)
```

<!-- grade_id: cell-dc1edc8aa056f275, points: 0.0, grade: true, solution: true -->
[YOUR ANSWER HERE]

```python

```

```python

```

```python

```